# Assignment 4

CSCI 1360: Foundations for Informatics and Analytics

## Important Dates

 - Deadline: 2022-07-13 at 5:00 pm
 
## Grading Breakdown

 - Q: 100 pts



# Overview

This assignment will focus on filesystem input/output.

If you have any questions or difficulties with the assignment, don't hesitate to send us an email!


## What to Submit
In this assignment, you will create one small python program. The name of the file should be YourLastName_A4.ipynb (for example Amirian_A4.ipynb). The source file (ending in .ipynb) should be submitted to the course eLC site. Detailed instructions on how to do this are included as the following part of this document.

## eLC Submission and Grading
After you have completed and thoroughly tested your program, upload and submit .ipynb file to eLC in the Assignment Dropbox. Always double check that your submission was successful on eLC and double check that you submitted the correct version of your completed program.  

The Assignment will be graded according to the following guidelines. 

•	A score between 0 and 100 will be assigned. 

•	If the source file(s) are not submitted before the specified deadline’s late period ends (48 hours after the deadline) or if they do not compile, then a grade of 0 will be assigned. 

•	25 points will be deducted for Assignments submitted late (note: Assignment can only be submitted up to 48 hours late, after 48 hours, a grade of zero will be assigned). 

# Question

In this question, we'll review the basics of file I/O (file input/output) and the various function calls and modes required.

### Part A

Write a function which:

 - is named `read_file_contents`
 - takes a string pathname as an argument
 - returns a single string that contains all the contents of the file
 
Don't import any additional packages.

If I have a file `random_text.txt`, I'll give the full path to this file to the function: `contents = read_file_contents("random_text.txt")`, and I should get back a single string `contents` that contains all the contents of the file.

**NOTE:** Your function should be able to handle errors gracefully! If an error occurs when trying to read from the file, your function should return `None` (note the capitalization of the first letter).

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def read_file_contents(filename):
    try:
        file = open(filename, 'r')
        content = ""
        
        for a in file.readlines():
            content += a
        file.close()
            
        return content
    except(FileNotFoundError):
        return None
    
print(read_file_contents("random_text.txt"))
```

Output:
```
None
```

<!-- grade_id: q2a_test1, points: 5, grade: true, solution: false -->
```python
truth = "This is some text.\nMore text, but on a different line!\nInsert your favorite meme here.\n"
pred = read_file_contents("A4/file1.txt")
assert truth == pred
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [13], in <cell line: 3>()
      1 truth = "This is some text.\nMore text, but on a different line!\nInsert your favorite meme here.\n"
      2 pred = read_file_contents("A4/file1.txt")
----> 3 assert truth == pred

AssertionError:
```

<!-- grade_id: q2a_test2, points: 5, grade: true, solution: false -->
```python
retval = -1
try:
    retval = read_file_contents("nonexistent/path.txt")
except:
    assert False
else:
    assert retval is None
```

### Part B

In this question, you'll read from one file, perform a simple computation, and write the results to a new file.

Write a function which:

 - is named `count_lines`
 - that takes two arguments: the first is a path to an input file to read, the second is the path to an output file to write
 - returns `True` on success, `False` if something went wrong (i.e an error occurred)
 
Your function will count the number of lines in the file at the first argument, and write this number to a file at the second argument.

**NOTE:** Your function should be able to handle errors gracefully! If an error occurs when trying to read from the file or write to the output file, your function should return `False`.

<!-- grade_id: q2b, points: 0, grade: false, solution: true -->
```python
def count_lines(input_file, output_file):
    lines = 0
    
    try:
        a = open(input_file, "r")
        while(a.readline() != ""):
            lines = lines + 1
            
        a.close()
        a = open(output_file, "w")
        a.write(str(lines))
        a.close()
        return True
    except:
        return False
```

<!-- grade_id: q2b_test1, points: 5, grade: true, solution: false -->
```python
import os
if os.path.exists("A4/file1_out.txt"):
    os.remove("A4/file1_out.txt")

assert count_lines("A4/file1.txt", "A4/file1_out.txt")
assert os.path.exists("A4/file1_out.txt")
assert int(open("A4/file1_out.txt", "r").read()) == 3
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [16], in <cell line: 5>()
      2 if os.path.exists("A4/file1_out.txt"):
      3     os.remove("A4/file1_out.txt")
----> 5 assert count_lines("A4/file1.txt", "A4/file1_out.txt")
      6 assert os.path.exists("A4/file1_out.txt")
      7 assert int(open("A4/file1_out.txt", "r").read()) == 3

AssertionError:
```

<!-- grade_id: q2b_test2, points: 5, grade: true, solution: false -->
```python
if os.path.exists("meaningless"):
    os.remove("meaningless")

r1 = None
try:
    r1 = count_lines("yet/another/nonexistent/path.txt", "meaningless")
except:
    assert False
else:
    assert not r1
```

<!-- grade_id: q2b_test3, points: 5, grade: true, solution: false -->
```python
r2 = None
try:
    r2 = count_lines("A4/file1.txt", "/this/should/throw/an/error.txt")
except:
    assert False
else:
    assert not r2
```

### Part C

In this question, you'll write a function which

 - is named `acount_lines` 
 - takes the same two arguments as `count_lines` in Part B
 - returns the same values as `count_lines` in Part B
 
This function performs the same operation as before, except in the case that the output file already exists: in this case, you'll **append** the line count to the file instead of overwriting it, thus preserving any existing previous line counts.

Each new appended line count should be *on its own line in the output file*. You may need to manually insert newline characters, which are a backslash followed by the letter n: **`\n`**

Your function should return `True` on success, and `False` if an error occurred.

**NOTE:** Your function should be able to handle errors gracefully! If an error occurs when trying to read from the file or write to the output file, your function should return `False`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
def acount_lines(inp, out):
    try:
        input1 = open(inp, "r")
        amountLines = len(input1.readlines())
        
        output1 = open(output, "a")
        output1.write(str(amountLines) + "\n")
        
        input1.close()
        output1.close()
        return True
    
    except:
        return False
    
acount_lines("input.txt", "output.txt")
```

Output:
```
False
```

<!-- grade_id: q2c_test1, points: 5, grade: true, solution: false -->
```python
if os.path.exists("A4/out_again.txt"):
    os.remove("A4/out_again.txt")

assert acount_lines("A4/file1.txt", "A4/out_again.txt")
assert os.path.exists("A4/out_again.txt")
assert int("".join(open("A4/out_again.txt", "r").read().split("\n"))) == 35
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [27], in <cell line: 4>()
      1 if os.path.exists("A4/out_again.txt"):
      2     os.remove("A4/out_again.txt")
----> 4 assert acount_lines("A4/file1.txt", "A4/out_again.txt")
      5 assert os.path.exists("A4/out_again.txt")
      6 assert int("".join(open("A4/out_again.txt", "r").read().split("\n"))) == 35

AssertionError:
```

<!-- grade_id: q2c_test2, points: 5, grade: true, solution: false -->
```python
assert acount_lines("A4/file2.txt", "A4/out_again.txt")
assert os.path.exists("A4/out_again.txt")
assert int("".join(open("A4/out_again.txt", "r").read().split("\n"))) == 35
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [28], in <cell line: 1>()
----> 1 assert acount_lines("A4/file2.txt", "A4/out_again.txt")
      2 assert os.path.exists("A4/out_again.txt")
      3 assert int("".join(open("A4/out_again.txt", "r").read().split("\n"))) == 35

AssertionError:
```

<!-- grade_id: q2c_test3, points: 5, grade: true, solution: false -->
```python
r1 = None
try:
    r1 = acount_lines("yet/another/nonexistent/path.txt", "meaningless")
except:
    assert False
else:
    assert not r1

r2 = None
try:
    r2 = acount_lines("A4/file2.txt", "/this/should/throw/an/error.txt")
except:
    assert False
else:
    assert not r2
```

```python

```